# Homework 4: Dictionaries, Functional Programming and Python Objects (24 pts)

name: Bryan Li

email: libryan@umich.edu

This homework assignment took me 7 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

I did not discuss this homework with anyone.

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

---
##  Problem 1: List Comprehensions (1 points)
Write a list comprehension where each element equals $3^{n}−1$ for $n = 1, 2, 3,..., 20$.

For ease of grading, please assign this list comprehension to a variable called `pow3minus1`.

<!-- grade_id: pow3minus1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pow3minus1 = [((3**i)-1)for i in range(1, 21)]
```

<!-- grade_id: correct_pow3minus1, points: 3, grade: true, solution: false -->
```python
assert pow3minus1[0:2] == [2, 8]
```

---
##  Problem 2: Iterator (2 points)
Deﬁne a class `Fibo` of iterators that enumerate the Fibonacci numbers. For the purposes of this problem, the Fibonacci sequence begins $0, 1, 1, 2, 3, ...$ with the $n$-th Fibonacci number $F_{n}$ given by the recursive formula $F_{n} = F_{n-1} + F_{n-2}$  where $F_{0} = 0$ and $F_{1} = 1$. 

Your solution should not make use of any function aside from addition (i.e., you should not need to use the function `fibo()` deﬁned in lecture a few weeks ago). Your class should support, at a minimum, an initialization method, a `__iter__` method (so that we can get an iterator) and a `__next__` method. 

**Note:** there is an especially simple solution to this problem that can be expressed in just a few lines using tuple assignment.

<!-- grade_id: Fibo, points: 0, grade: false, solution: true -->
```python
class Fibo():
    # YOUR CODE HERE
    def __init__(self):
        self.F0 = 0
        self.Fnm2 = 0
        self.F1 = 1
        self.Fnm1 = 1
        self.i = 0
        
    def __next__(self):
        if self.i == 0:
            self.i += 1
            return(self.F0)
        elif self.i == 1:
            self.i += 1
            return(self.F1)
        else:
            self.n = self.Fnm1 + self.Fnm2
            self.Fnm2 = self.Fnm1
            self.Fnm1 = self.n     
            return(self.n)

    def __iter__(self):
        return(self)
    
```

<!-- grade_id: correct_fibo, points: 2, grade: true, solution: false -->
```python
ans1 = Fibo()
assert [next(ans1) for _ in range(10)] == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34]
```

##  Problem 3: Dictionaires (9 points)

###  3.1: Find the major element (3 points)
Given an array nums of size n, return the majority element.

The majority element is the element that appears more than ⌊n / 2⌋ times. You may assume that the majority element always exists in the array.

<!-- grade_id: find_major_element, points: 0, grade: false, solution: true -->
```python
import numpy as np

def find_major_element(nums):
    # YOUR CODE HERE
    nums = np.array(nums)
    uniq_nums = []
    for num in nums:
        if num not in uniq_nums:
            uniq_nums.append(num)

    for num in uniq_nums:
        if len(nums[nums == num]) > len(nums)/2:
            return(num)
```

<!-- grade_id: correct_find_major_element, points: 3, grade: true, solution: false -->
```python
assert find_major_element([1]) == 1
assert find_major_element([-1, 1, 1, 1, 1, -1, -1]) == 1
```

###  3.2: Group Anagram (3 points)
Given an array of strings, group the anagrams together. You can return the answer in any order.

An Anagram is a word or phrase formed by rearranging the letters of a different word or phrase, using all the original letters exactly once.

<!-- grade_id: group_anagrams, points: 0, grade: false, solution: true -->
```python
import numpy as np

def group_anagrams(my_string):
    # YOUR CODE HERE
    return_l = []
    my_string = np.array(my_string)

    for string in my_string:
        for l in return_l:
            if string in return_l:
                break
        # Get counts of each letter in current string
        components = []
        for let in string:
            components.append(let)
        components = np.array(components)
        uniq_lets = []
        for let in string:
            if let not in uniq_lets:
                uniq_lets.append(let)
        let_dict = {}
        for let in uniq_lets:
            let_dict[let] = len(components[components == let])

        anagrams = [string]
        for string2 in my_string[my_string != string]:            
            # Get counts of each letter in string
            components2 = []
            for let in string2:
                components2.append(let)
            components2 = np.array(components2)
            uniq_lets2 = []
            for let in string2:
                if let not in uniq_lets2:
                    uniq_lets2.append(let)
            let_dict2 = {}
            for let in uniq_lets2:
                let_dict2[let] = len(components2[components2 == let])

            if let_dict == let_dict2:
                anagrams.append(string2)
            
        return_l.append(anagrams)

        return_l2 = []
        for l in return_l:
            l.sort()
            if l not in return_l2:
                return_l2.append(l)
                
    return(return_l2)
```

```python
print(group_anagrams(["", "b"]))
print(group_anagrams(['listen', 'silent', 'enlist']))
print(group_anagrams(["eat","tea","tan","ate","nat","bat"]))
print(group_anagrams(["Tea", "Eat", "tea", "eat"]))
```

Output:
```
[[''], ['b']]
[['enlist', 'listen', 'silent']]
[['ate', 'eat', 'tea'], ['nat', 'tan'], ['bat']]
[['Tea'], ['Eat'], ['eat', 'tea']]
```

<!-- grade_id: correct_group_anagrams, points: 3, grade: true, solution: false -->
```python
assert group_anagrams(["", "b"]) == [[''], ['b']]
assert group_anagrams(["listen", "silent", "enlist"]) == [['listen', 'silent', 'enlist']]
assert group_anagrams(["eat","tea","tan","ate","nat","bat"]) == [['eat', 'tea', 'ate'], ['tan', 'nat'], ['bat']]
assert group_anagrams(["Tea", "Eat", "tea", "eat"]) == [['Tea'], ['Eat'], ['tea', 'eat']]
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[134], line 2
      1 assert group_anagrams(["", "b"]) == [[''], ['b']]
----> 2 assert group_anagrams(["listen", "silent", "enlist"]) == [['listen', 'silent', 'enlist']]
      3 assert group_anagrams(["eat","tea","tan","ate","nat","bat"]) == [['eat', 'tea', 'ate'], ['tan', 'nat'], ['bat']]
      4 assert group_anagrams(["Tea", "Eat", "tea", "eat"]) == [['Tea'], ['Eat'], ['tea', 'eat']]

AssertionError:
```

###  3.3: Contain Duplicate (3 points)
Given an integer array nums (represented as Python list) and an integer k, return true if there are two distinct indices i and j in the array such that nums[i] == nums[j] and abs(i - j) <= k.



<!-- grade_id: contain_duplicate, points: 0, grade: false, solution: true -->
```python
def contain_duplicate(nums, k):
    # YOUR CODE HERE
    uniq_nums = []
    for num in nums:
        if num not in uniq_nums:
            uniq_nums.append(num)

    count_dict = {}
    i_dict = {}
    for uniq_num in uniq_nums:
        i = 0
        i_l = []
        count = 0
        for num in nums:
            if uniq_num == num:
                i_l.append(i) 
                count += 1
                i += 1
            else:
                i += 1
        count_dict[uniq_num] = count
        i_dict[uniq_num] = i_l

    # Check for pair of letters
    num_pair = []
    for num in count_dict:
        if count_dict[num] == 2:
            num_pair.append(num)
    if len(num_pair) < 1:
        return False

    # Check for k condition
    for num in num_pair:
        if abs(i_dict[num][0] - i_dict[num][1]) <= k:
            return True

    return False
```

<!-- grade_id: correct_duplicate, points: 3, grade: true, solution: false -->
```python
assert contain_duplicate([1, 2, 3, 1], 3) == True
assert contain_duplicate([], 0) == False
assert contain_duplicate([1, 2, 3, 4, 5, 6, 7, 8, 9, 10], 5) == False
```

## Problem 4: Vectors with an object-oriented approach.(12 points)
**Note:** Before starting this problem, you are strongly encouraged to read _Chapter 14: Python Objects_ in the textbook: https://www.py4e.com/html3/14-objects

<!-- grade_id: vector, points: 0, grade: false, solution: true -->
```python
"""Include all the code for the Vector class here."""
import math
import numpy as np

class Vector:
    # YOUR CODE HERE
    def __init__(self, dim=0, nums=None):
        if dim == 0:
            self.dimension = 0
            self.entries = tuple()
        elif type(nums) != list and type(nums) != tuple:
            raise TypeError('Entries must be in a list or tuple')
        elif dim < 0:
            raise ValueError('Invalid dimension')
        elif dim != len(nums):
            raise ValueError('Dimension and values do not match')
        else:
            self.dimension = dim
            self.entries = tuple(nums)

    def __eq__(self, V2):
        return self.entries == V2.entries
    def __gt__(self, V2):
        return self.entries > V2.entries
    def __le__(self, V2):
        return self.entries < V2.entries

    def dot(self, V2):
        dot_prod = 0
        for i in range(0, len(self.entries)):
            dot_prod += (self.entries[i] * V2.entries[i])
        return(dot_prod)

    def __mul__(self, V2):
        if type(V2) == float or type(V2) == int:
            return Vector(len(self.entries), tuple(np.array(self.entries) * V2))
        elif type(V2) == Vector:
            if self.dimension != V2.dimension:
                raise ValueError('Incompatible vector dimensions')
            else:
                new_entries = []
                for i in range(0, len(self.entries)):
                    new_entries.append(self.entries[i] * V2.entries[i])
                return(Vector(len(new_entries), new_entries))
    def __rmul__(self, V2):
        return self * V2

    def norm(self, p):
        if p == 0:
            ary = np.array(self.entries)
            return len(ary[ary != 0])
        elif p < math.inf:
            ary = np.array(self.entries)
            return sum(abs(ary)**p)**(1/p)
        else:
            ary = np.array(self.entries)
            return max(ary)
```

### Part 1 (5 points)
Deﬁne a class `Vector`. The initializer for your class should take the dimension as its ﬁrst argument and a `list` or `tuple` of numbers (`int`s or `ﬂoat`s), representing the vector’s entries, as its second argument (**Note:** your initializer should work correctly given either a list or a tuple of numbers). The initializer should create a vector with dimension $0$ and no entries (an empty tuple `tuple()`) if the user does not provide the corresponding arguments.

Your class vector should have two __instance__ attributes: `dimension` and `entries`. The first containing the dimension of the vector as an `int` and the second containing the entries of the vector as a `tuple`. 

The initializer should raise an appropriate error in the case where the dimension is invalid (i.e., wrong type or a negative number), and should also raise an error in the event that the dimension and the number of supplied entries disagree.

**Note:** Please raise `ValueError` or `TypeError` for your error checking in this problem. Raising either will work. 

<!-- grade_id: vec1, points: 1, grade: true, solution: false -->
```python
"""Check initialization works on sensible cases"""
v = Vector(4, (1, 2, 3, 4))
v = Vector(4, [1, 2, 3, 4])
v = Vector(1, [1])
v = Vector(2, [2, 3])
v = Vector(1, (1,))
v = Vector(0, [])
v = Vector(0)
assert True
```

<!-- grade_id: vec2, points: 1, grade: true, solution: false -->
```python
"""Vector has correct attributes"""
v = Vector(4, (1, 2, 3, 4))
assert isinstance(v, Vector)
assert 'entries' in v.__dict__.keys()
assert 'dimension' in v.__dict__.keys()
assert type(v.dimension) == int
assert type(v.entries) == tuple
```

<!-- grade_id: vec3, points: 1, grade: true, solution: false -->
```python
"""Check you can get entries from vector and they are correct when retrieved"""
v = Vector(3, [1, 2, 3])
assert tuple(v.entries) == (1, 2, 3)
```

<!-- grade_id: vec4, points: 2, grade: true, solution: false -->
```python
"""Check exceptions are handled correctly"""
try:
    Vector(5, (1, 2, 3, 4))
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector(-4, (1, 2, 3, 4))
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector('a', [1, 2, 3, 4])
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector(1, [])
    assert False
except (ValueError, TypeError):
    assert True
```

### Part 2 (2 point)
Implement the necessary operator(s) to support comparison (equality, less than, less or equal to, greater than, etc) of `Vector` objects. We will say that two `Vector` objects are equivalent if they have the same coordinates. Otherwise, comparison should be analogous to `tuple`s in Python, so that comparison is done on the ﬁrst coordinate ﬁrst, then the second coordinate, then the third, and so on. So, for example, the two-dimensional vector `(2, 4)` is ordered before (less than) `(2, 5)`. 

You can assume you are always comparing vectors of the same dimension. 

<!-- grade_id: vecpart2, points: 2, grade: true, solution: false -->
```python
assert Vector(3, [2, 1, 1]) != Vector(3, [1, 3, 3])
assert Vector(3, [2, 1, 1]) >= Vector(3, [1, 3, 3])
assert Vector(3, [1, 1, 1]) == Vector(3, [1, 1, 1])
assert Vector(3, [1, 3, 3]) <= Vector(3, [2, 1, 1])
```

### Part 3 (1 point) 
Implement a method `Vector.dot` that takes a single Vector as its argument and returns the inner product of the caller with the given `Vector` object. You can assume the input is always a vector of the same dimension as the original one.

<!-- grade_id: vecpart3, points: 1, grade: true, solution: false -->
```python
"""Check your Vector.dot() method works correctly"""
assert Vector(3, (1, 2, 3)).dot(Vector(3, (1, 2, 2))) == 11.0
```

### Part 4 (2 points)
We would also like our Vector class to support scalar multiplication. Left- or right multiplication by a scalar, e.g., `2*v` or `v*2`, where `v` is a Vector object, should result in a new `Vector` object with its entries all scaled by the given scalar. We will also follow `R` and `numpy` (which you will learn in a few weeks), and use `*` to denote entrywise vector-vector multiplication, so that for `Vector` objects `v` and `w`, `v*w` results in a new `Vector` object, with the i-th entry of `v*w` equal to the i-th entry of `v` multiplied by the i-th entry of `w`. Implement the appropriate operators to support this multiplication operation. 

Many languages have a convention for dealing with multiplication of vectors that diﬀer in their dimension, but we will not handle this case. Your method should raise a `ValueError` in the event that `v` and `w` disagree in their dimensions.

<!-- grade_id: partvec41, points: 1, grade: true, solution: false -->
```python
assert 2 * Vector(3, (1, 2, 3)) == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * 2 == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * Vector(3, (1, 2, 3)) == Vector(3, (1, 4, 9))
```

<!-- grade_id: vecpart42, points: 1, grade: true, solution: false -->
```python
try:
    Vector(3, (1, 2, 3)) * Vector(2, (1, 2))
    assert False
except ValueError:
    assert True
```

### Part 5 (2 points)
For a real number $0 \leq p \leq \infty$, and a vector $v \in \mathbb{R}^{d}$, the $p$-norm of $v$, written $\| v \|_{p}$, is given by: 

$$ 
\| v \|_p = \begin{cases}
                        \sum_{i=1}^d 1_{v_i \neq 0} &\mbox{ if } p = 0\\
                        (\sum_{i=1}^d |v_i|^p )^{1/p} &\mbox{ if }
                                                        0 < p < \infty, \\
                        \max_{i=1,2,\dots,d} |v_i| &\mbox{ if } p=\infty
                        \end{cases}
$$

where $v_i$ is the $i$-th component of the vector and the notation $1_{v_i \neq 0}$ indicates: 

$$ 
1_{v_i \neq 0} = \begin{cases}
                        1 &\mbox{ if } v_i \neq 0\\
                        0 &\mbox{ if } v_i = 0
                        \end{cases}. 
$$


Implement a method `Vector.norm` that takes a single int or ﬂoat `p` as an argument and returns the $p$-norm of the calling `Vector` object. You can assume the input $p$ is always a non-negative integer or float.

`Hint:` see https://docs.python.org/3/library/functions.html# float for documentation on representing positive inﬁnity in Python.

<!-- grade_id: vecpart5, points: 1, grade: true, solution: false -->
```python
"""Check your Vector.dot() method works correctly"""
assert Vector(3, (1, 2, 3)).norm(1) == 6.0
assert Vector(2, (1, 2)).norm(2) == math.sqrt(1**2 + 2**2)
```

```python

```